## Week 6 Assignment - W200 Python for Data Science, UC Berkeley MIDS

Write code in this jupyter Notebook to solve each of the following problems. Each problem should have its solution in a separate cell. Please upload this **Notebook**, your **scrabble.py** file, the **sowpods.txt** file and your **score_word** module with your solutions to your GitHub repository in your SUBMISSIONS/week_06 folder by 11:59PM PST the night before class.

## Objectives:

- Read and understand PEP 8 standards
- Use all of your previously gained knowledge together on a single program
- Demonstrate how to import a user made module and function into python from another .py file
- Demonstrate how to input command line arguments into a .py file

## 6-1. PEP 8 Style Guide (reading and response)

Your first task for this week is to write a **250** word reading response to the article below. In addition, please list **3 questions** that you have from the article. Please write your response in a markdown cell, below this cell in the notebook.

The writing response is a free response, so you may write about your reactions. An interesting thing that you saw in the article, something that really stuck out to you, etc.

**Article**: [The PEP 8 Style Guide](https://www.python.org/dev/peps/pep-0008). This document is really important for Python coders because it describes best practices and customs for how one should write Python code. Please read it all and prepare your questions for class.

<!-- grade_id: cell-eaae902b022d20da, points: 20, grade: true, solution: true -->
#### Alex West

As a complete beginner to programming in general, "readability" was initially a difficult concept. As I've become conversational, however, it makes perfect sense to me that style guides on Python are meticulously maintained. "Code is read much more often than it is written" -- this stuck out to me right off the bat, since I hadn't really thought about it except when I do my own homework. But it makes sense, and because of this, code needs to be readable by a diverse group of people. Hence, agreed upon style guidelines. 

Overall, what stuck out to me was the unbelievable detail. Again, upon reflection, not surprising. Of course there should be rules about 4 spaces vs. tabs, 79 characters per line, etc. It's not unlike writing grammar rules for a foreign language. However, the "2 spaces after a sentence ending period" within a comment rule seems excessive and a product of a bygone era. The font is reminiscent of typewriters, so maybe that's where it comes from. But it certainly isn't followed in contemporary written word. In addition, the rule that non-English speakers should write comments in English unless "120% sure" it will never be read by anyone outside the language... seemed a bit dramatic and out of character for the authors of a programming language. 

I learned a lot reading through the style guide. Much of what was there made sense, and it showed me that the human brain has a natural capacity for form and style. In other words, I was internalizing a lot of this style convention without knowing it, simply by coding along with the professors during the async material. Looking at some of the "No" examples I automatically shook my head; "It looks wrong." How did I know that? 

As with most grammar and style, reading and internalizing the conventions makes you a better writer / coder. For example, "When catching exceptions, mention specific exceptions whenever possible instead of using a bare except: clause." This resonated in particular. If you keep you exceptions specific, you'll be writing better code. 

There were elements to the style guide I couldn't quite understand (the API stuff specifically) but it's nice to know that it's there whenever I need to look something up. 


## 6-2. Cheating at Scrabble

Write a Python script that takes a Scrabble rack as a command-line argument and prints all valid Scrabble words that can be constructed from that rack, along with their Scrabble scores, sorted by score. Valid Scrabble words are provided in the data source below. A Scrabble rack is made up of a maximum of any 7 characters.

Below are the requirements for the program:
- This needs to be able to be run as a command line tool as shown below (not an input statement!)
- Please name the python file: `scrabble.py`
- Allow anywhere from 2-7 character tiles to be inputted 
- Do not worry about the number of the same tiles (e.g. a user is allowed to input ZZZZZQQ)
- Output the **total** list of words as (score, word) tuples, sorted by the score as shown below
- Output at the end the 'Total number of words:' that can be made with the letters
- Please include a function called `score_word` in a separate module. Import this function into your main solution code.
- You need to handle input errors from the user and suggest what that error might be caused by and how to fix it (i.e. a helpful error message)
- Implement wildcards as either `*` or `?`. That is, let the user specify a wildcard character that can take any value. There can be a total of two wild cards in any user input (one of each character). Only use the `*` and `?` as wildcard characters
- Wildcard characters are scored as 0 points, just like in the real Scrabble game
- Your program should take less than a minute to run with 2 wildcards in the input - if it is more time than that your algorithm is not optimized very well!
- Write docstrings for the functions and puts comments in your code.

Extra Credit (+10 points):
Allow a user to specify that a certain letter has to be at a certain location. Your program must work without it so this is completely optional. For the extra credit, locations of certain letters must be specified at the command line, it may not be some sort of user prompt.  (Please put a sample of how to run your extra credit & comments in the extra credit cell of this notebook below!)

An example invocation and output:
```
$ python scrabble.py ZAEFIEE
(17, feeze)
(17, feaze)
(16, faze)
(15, fiz)
(15, fez)
(12, zee)
(12, zea)
(11, za)
(6, fie)
(6, fee)
(6, fae)
(5, if)
(5, fe)
(5, fa)
(5, ef)
(2, ee)
(2, ea)
(2, ai)
(2, ae)
Total number of words: 19
```

The Data
http://courses.cms.caltech.edu/cs11/material/advjava/lab1/sowpods.zip contains all words in the official SOWPODS word list, one word per line. You should download the word file and keep it in your repository so that the program is standalone (instead of accessing it over the web from Python).

You can read data from a text file with the following code:

```
with open("sowpods.txt","r") as infile:
    raw_input = infile.readlines()
    data = [datum.strip('\n') for datum in raw_input]
```
This will show the first 6 words:
```
print(data[0:6])
```

Please use the dictionary below containing the letters and their Scrabble values:

```
scores = {"a": 1, "c": 3, "b": 3, "e": 1, "d": 2, "g": 2,
         "f": 4, "i": 1, "h": 4, "k": 5, "j": 8, "m": 3,
         "l": 1, "o": 1, "n": 1, "q": 10, "p": 3, "s": 1,
         "r": 1, "u": 1, "t": 1, "w": 4, "v": 4, "y": 4,
         "x": 8, "z": 10}
```

Tips:
- We recommend that you work on this and try to break down the problems into steps on your own before writing any code. Once you've scoped generally what you want to do, start writing some code and if you get stuck, take a step back and go back to thinking about the problem rather than trying to fix lots of errors at the code level. You should only use the Python standard library in this assignment, however any tool in the standard library is fair game.
- If you keep getting stuck check out: https://openhatch.org/wiki/Scrabble_challenge. This is where we got the idea for this assignment and it provides some helpful tips for guiding you along the way. If that link doesn't work you can use the Google cached version here. However, we would recommend that you try to implement this first before looking at the hints on the website.

Good luck!

### The code below will test your command line implementation of the scrabble.py code. We've made some of these tests are available for you to try!

```python
# Code for the testing

import subprocess
from nose.tools import assert_equal 
from nose.tools import assert_true
from nose.tools import assert_greater
from nose.tools import assert_less
```

<!-- grade_id: cell-f55d7e98eba40a10, points: 20, grade: true, solution: false -->
```python
""" Checks that the code runs and checks one user error messages 
    (this is just one of many user errors you should be checking for!)
"""

!python scrabble.py  # no rack error


```

<!-- grade_id: cell-1893c284e867334a, points: 10, grade: true, solution: false -->
```python
""" Does not fail due to trivial mistakes and takes correct wildcard characters """

!python scrabble.py PENguin    # does not fail due to case
!python scrabble.py PEN*?in    # takes wildcards
```

<!-- grade_id: cell-7f1c7fe54a38a19e, points: 5, grade: true, solution: false -->
```python
""" The code should produce a list of all words from the rack with scores """

#test the code in the command line
!python scrabble.py PENGUIN

```

<!-- grade_id: cell-bb7a872eda80e6a5, points: 5, grade: true, solution: false -->
```python
# Autograding test
```

<!-- grade_id: cell-8d2de0735f148301, points: 5, grade: true, solution: false -->
```python
""" The code should produce a list of all words from the rack with scores """
!python scrabble.py PENGU*?

```

<!-- grade_id: cell-16d4589caf032f51, points: 5, grade: true, solution: false -->
```python
# Autograding test
```

<!-- grade_id: cell-ceab3d4433043319, points: 10, grade: true, solution: false -->
```python
""" The code should run in seconds to a few minutes """
import time
start=time.time()

#test the code in the command line
cmd = [ 'python', 'scrabble.py', 'PENGU*?' ]
out=bytes.decode(subprocess.Popen( cmd, stdout=subprocess.PIPE ).communicate()[0])

tot_time=time.time()-start
print('Total time was {} seconds'.format(tot_time))
assert_less(tot_time, 300)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-1-e6f723e52cf7> in <module>()
      5 #test the code in the command line
      6 cmd = [ 'python', 'scrabble.py', 'PENGU*?' ]
----> 7 out=bytes.decode(subprocess.Popen( cmd, stdout=subprocess.PIPE ).communicate()[0])
      8 
      9 tot_time=time.time()-start

NameError: name 'subprocess' is not defined
```

<!-- grade_id: cell-157fccf420d13535, points: 10, grade: true, solution: true -->
```python
"""Implement extra credit call run the code here
   If this cell isnt filled out - we'll assume the extra credit wasn't done
   Please write a comment on how to use your extra credit syntax also!
"""
# YOUR EXTRA CREDIT CODE HERE
```

## If you have feedback for this homework, please submit it using the link below:

http://goo.gl/forms/74yCiQTf6k